# Assignment 3: Neuroscience Literature Mining

This assignment will demonstrate how to pull data from an API, how to work with the LISC package, and test your knowledge of working with classes and numpy.

This assignment is worth 5 points (5% of your grade). 

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INCLUDED IN THE ASSIGNMENT.**


## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (meaning, delete) these lines of code with code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.


## Tips & Tricks

The following are a couple tips & tricks that may help you if you get stuck on anything.

#### Printing Variables
You can (and should) print and check variables as you go. This allows you to check what values they hold, and fix things if anything unexpected happens.

#### Restarting the Kernel
- If you run cells out of order, you can end up overwriting things in your namespace. 
- If things seem to go weird, a good first step is to restart the kernel, which you can do from the kernel menu above.
- Even if everything seems to be working, it's a nice check to 'Restart & Run All', to make sure everything runs properly in order.

# Part 1: Fetching data from URLs

In this section we'll show you the very basics of how to look up how many articles are indexed in Pubmed that have specific words or phrases (n-grams) in their titles and/or abstracts.

For this example, we begin by counting how many articles have been published whose titles or abstracts contain the phrases "working memory" or "short term memory". Look through the cell below to get a sense of what it is doing, and then run the cell.


```python
import urllib.request # default library for requesting data from URLs

# this is the base URL for making use of the NIH NLM database search API
u_eutils = 'https://eutils.ncbi.nlm.nih.gov/entrez/eutils/esearch.fcgi'

##### all of the code that follows is to construct the full search URL #####

# what database to search; here, pubmed
u_db = 'db=pubmed'

# what information to return; here, count of the number of articles
u_rettype = 'rettype=count'

# what field to search; here, tiab (pubmed's code for TItle and ABstract)
u_field = 'field=tiab'

# what term(s) to search for; here "working memory" OR "short term memory"
# note that the double quotes around the phrases indicate we only want to return
    # searches that contain that exact phrase
u_term = 'term='
url_searchterms = '"working memory" OR "short term memory"'
url_searchterms = url_searchterms.replace(' ', '+') # URLs don't do spaces; replace with +
u_term = u_term + url_searchterms

# stitch the URL together from the parts we gave it
url = u_eutils + '?' + u_db + '&' + u_rettype + '&' + u_field + '&' + u_term

##### end URL construction block #####

with urllib.request.urlopen(url) as response:
    xml = response.read()
xml
```

Output:
```
b'<?xml version="1.0" encoding="UTF-8" ?>\n<!DOCTYPE eSearchResult PUBLIC "-//NLM//DTD esearch 20060628//EN" "https://eutils.ncbi.nlm.nih.gov/eutils/dtd/20060628/esearch.dtd">\n<eSearchResult>\n\t<Count>38800</Count>\n</eSearchResult>\n'
```

## Q1: Making your own search

You can see, between the `<Count>` tags in the long response above, that this returns (on January 30th, 2020) 38763 articles in Pubmed with the phrases "working memory" or "short term memory" in their titles and/or abstracts. This number, of course, will keep increasing over time (and might have increased since we released the assignment!).

1. Use the eutils documentation, specifically the esearch tools [here](https://www.ncbi.nlm.nih.gov/books/NBK25499/#chapter4.ESearch) and [here](https://www.ncbi.nlm.nih.gov/books/NBK3827/#pubmedhelp.Title_TI), to modify the above code to find the count of how many articles have been published with the following three terms <font color='red'>**in their titles only**</font>:

    * dentate gyrus
    * entorhinal cortex
    * subiculum


2. In the cell below the url creation cell (after you've retrieved the xml) create a new variable, `number_of_articles`, that lists that count (as an integer). Remember the double quotes around the phrases.

<!-- grade_id: cell-de51725c9709625a, points: 0, grade: false, solution: true -->
```python
# Run API query here

import urllib.request # default library for requesting data from URLs

# this is the base URL for making use of the NIH NLM database search API
u_eutils = 'https://eutils.ncbi.nlm.nih.gov/entrez/eutils/esearch.fcgi'

##### all of the code that follows is to construct the full search URL #####

# what database to search; here, pubmed
u_db = 'db=pubmed'

# what information to return; here, count of the number of articles
u_rettype = 'rettype=count'

# what field to search; here, tiab (pubmed's code for TItle and ABstract)
u_field = 'field=ti'

# what term(s) to search for; here "working memory" OR "short term memory"
# note that the double quotes around the phrases indicate we only want to return
    # searches that contain that exact phrase
u_term = 'term='
url_searchterms = '"dentate gyrus" OR "entorhinal cortex" OR "subiculum"'
url_searchterms = url_searchterms.replace(' ', '+') # URLs don't do spaces; replace with +
u_term = u_term + url_searchterms

# stitch the URL together from the parts we gave it
url = u_eutils + '?' + u_db + '&' + u_rettype + '&' + u_field + '&' + u_term

##### end URL construction block #####

with urllib.request.urlopen(url) as response:
    xml = response.read()
xml
```

Output:
```
b'<?xml version="1.0" encoding="UTF-8" ?>\n<!DOCTYPE eSearchResult PUBLIC "-//NLM//DTD esearch 20060628//EN" "https://eutils.ncbi.nlm.nih.gov/eutils/dtd/20060628/esearch.dtd">\n<eSearchResult>\n\t<Count>5095</Count>\n</eSearchResult>\n'
```

<!-- grade_id: cell-3963a9249af0b851, points: 0, grade: false, solution: true -->
```python
# Create number_of_articles here
number_of_articles = 5095
```

<!-- grade_id: cell-54c47de7ad67b3af, points: 5, grade: true, solution: false -->
```python
# Tests for Q1 (worth 5 points)
assert isinstance(u_field,str)
assert isinstance(url_searchterms,str)
assert isinstance(url,str)
assert isinstance(number_of_articles,int)
```

<!-- grade_id: cell-96efe836ba4d28aa, points: 5, grade: true, solution: false -->
```python
# Hidden tests for Q1 (worth 5 points)
```

<!-- grade_id: cell-91e9689f112f83dc, points: 5, grade: true, solution: false -->
```python
# Hidden tests for Q1 (worth 5 points)
```

<!-- grade_id: cell-7f45cca832af3943, points: 5, grade: true, solution: false -->
```python
# Additional hidden tests for Q1 (worth 5 points)
```

# Part 2: Basic literature scanner (LISC) usage

Now, instead of all of this manual URL construction, we're going to take advanatage of an open-source toolbox from Tom Donoghue and the Voytek lab here at UC San Diego called LISC (for LIterature SCanner). It's available for pip install from pypi [here](https://pypi.org/project/lisc/).

You can install it on your own system by opening your terminal, going into your Pyton environment, and typing the following:

<code>pip install lisc</code>

Or, you can install it in this Jupyter notebook by running the following cell:

```python
! pip install --user lisc
```

Output:
```
Requirement already satisfied: lisc in /home/adfrankl/.local/lib/python3.6/site-packages (0.1.1)
Requirement already satisfied: lxml in /opt/conda/lib/python3.6/site-packages (from lisc) (4.4.1)
Requirement already satisfied: nltk in /home/adfrankl/.local/lib/python3.6/site-packages (from lisc) (3.4.5)
Requirement already satisfied: beautifulsoup4 in /opt/conda/lib/python3.6/site-packages (from lisc) (4.6.3)
Requirement already satisfied: requests in /opt/conda/lib/python3.6/site-packages (from lisc) (2.12.4)
Requirement already satisfied: numpy in /opt/conda/lib/python3.6/site-packages (from lisc) (1.17.3)
Requirement already satisfied: six in /opt/conda/lib/python3.6/site-packages (from nltk->lisc) (1.12.0)
```

## Demonstration of LISC
The below cell is all of the LISC imports you'll need.

Here, we create a list called <code>terms_a</code>, that contains the same three search terms we used above:

* dentate gyrus
* entorhinal cortex
* subiculum

Note the double quotes around the phrases. This time we're passing an argument for field that searches both titles and abstracts (<code>field='tiab'</code>).

```python
from lisc import Counts
from lisc.utils.db import SCDB
from lisc.plts.counts import *

terms_a = [['"dentate gyrus"'], ['"entorhinal cortex"'], ['"subiculum"']]

# Initialize counts object and use the add_terms method to add terms that we want to search
counts = Counts()
counts.add_terms(terms_a)

# Collect data using the run_collection method
counts.run_collection(verbose=True, db='pubmed', field='tiab')

# Check how many articles were found for each search term
counts.check_counts()
```

## Q2
`check_counts` is a(n):

* `A` attribute of `counts`
* `B` method of `counts`
* `C` inherited class of `counts`

Write your answer below as a new variable, <code>q2_answer</code>. So if the answer was a hypothetical option <code>D</code>, you would write:
​
<code>q2_answer = 'D'</code>

<!-- grade_id: cell-43e81e8682b5c1b3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
q2_answer = 'B'
```

<!-- grade_id: cell-ade9472e1101f50f, points: 5, grade: true, solution: false -->
```python
# Tests for Q2 (worth 5 points)

assert isinstance(q2_answer,str)

```

## Q3
1. Create a new list called <code>terms_b</code>, that contains the four lobes of the brain: 'frontal lobe', 'temporal lobe', 'parietal lobe', and 'occipital lobe'. Remember the double quotes around the phrases.
2. Create a new variable, <code>number_of_lobes_articles</code>, that is equal to the *total* number of articles (as an integer) that mention any of those phrases in their titles or abstracts.

<!-- grade_id: cell-0cccd1557823b047, points: 0, grade: false, solution: true -->
```python
# Run your LISC query here
terms_b = [['"frontal lobe"'],['"temporal lobe"'],['"parietal lobe"'],['"occipital lobe"']]
# Initialize counts object and use the add_terms method to add terms that we want to search
counts = Counts()
counts.add_terms(terms_b)

# Collect data using the run_collection method
counts.run_collection(verbose=True, db='pubmed', field='tiab')

# Check how many articles were found for each search term
counts.check_counts()
```

<!-- grade_id: cell-cd33a4d9813aa339, points: 0, grade: false, solution: true -->
```python
# Assign number_of_lobes_articles here
number_of_lobes_articles = 63373
```

<!-- grade_id: cell-7c20250c131d47d4, points: 5, grade: true, solution: false -->
```python
# Tests for Q3, wroth 5 points
assert isinstance(counts,object)
assert isinstance(number_of_lobes_articles,int)
```

<!-- grade_id: cell-e84ad3a097c3031e, points: 5, grade: true, solution: false -->
```python
# Hidden tests, worth 5 points
```

## LISC: Term co-occurances
LISC not only finds the counts for each search time, but also runs each pairwise comparison between terms to find how many papers are published that talk about <code>term_i</code> <em>and</em> <code>term_j</code>. This is reported as an array in <code>counts.counts</code> where each <em>i,j</em> index is how many papers are published that mention <code>term_i</code> with <code>term_j</code>.

For example, for the four lobes, we get the following:

```python
counts.counts
```

The top row and left column are co-occurances for the the zeroth item in our list, "frontal lobe", with itself and all other terms. You'll note that LISC sets a term as having 0 publications matching with itself, by convention. You see that there are more papers published that mention "frontal lobe" with "parietal lobe" (>6500 together) than either of the other two lobes.

## Q4
Using LISC and the <code>terms_a</code> list, store the <code>counts.counts</code> array as a new variable, <code>cooccurance_matrix</code>.

<!-- grade_id: cell-b832df074d20ae47, points: 0, grade: false, solution: true -->
```python
# Initialize counts object and add the terms that we want to search
counts = Counts()
counts.add_terms(terms_a)

# Collect data using the run_collection method
counts.run_collection(verbose=True, db='pubmed', field='tiab')

cooccurance_matrix = counts.counts
cooccurance_matrix

plot_matrix(cooccurance_matrix, counts.terms['A'].labels, counts.terms['A'].labels)
```

<!-- grade_id: cell-0e62aed37f8722c7, points: 5, grade: true, solution: false -->
```python
# Tests for Q4 (worth 5 points)
import numpy as np
assert isinstance(cooccurance_matrix,np.ndarray)

```

## Q5
At this point, you might have noticed that the output of counts is an numpy array. Use the cell below to assign the **shape** of your numpy array to a variable `matrix_shape`.

<!-- grade_id: cell-8ef3fbc6ecff9b5e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
matrix_shape = cooccurance_matrix.shape
matrix_shape
```

<!-- grade_id: cell-5f0777c01bb269c8, points: 5, grade: true, solution: false -->
```python
# Tests for Q5, worth 5 points
assert isinstance(matrix_shape,tuple)

```

## Q6
Which pair has the most papers published talking about them together?

* <code>A</code>: dentate gyrus with entorhinal cortex
* <code>B</code>: dentate gyrus with subiculum
* <code>C</code>: entorhinal cortex with subiculum

Write your answer below as a new variable, <code>q6_answer</code>. So if the answer was a hypothetical option <code>D</code>, you would write:

<code>q6_answer = 'D'</code>

<!-- grade_id: cell-43e81e8682b5c1b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
q6_answer = 'A'
```

<!-- grade_id: cell-ade9472e1101f50e, points: 5, grade: true, solution: false -->
```python
# Tests for Q6 (worth 5 points)

assert isinstance(q6_answer,str)

```

```python

```